---
title: chap2
description: blah
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 72dc9c55a5
xp: 100
```

helllo lmkj lmkjmlkj k mlk jklj mlkjmlj ml jklmk jmlkjmlkjkjlmkj!!!

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
