---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise. In fact it's an awesome example exercise!!11!!11

woa

`@instructions`
It should be clear at this point that the final objective of this course is awesomeness.

With this in mind, just go ahead and print the word "awesome".

`@hint`
It might have something to do with the `print` function...

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}
print('awesome')
```

`@sct`
```{python}
Ex().has_output("awesome", pattern = False)
success_msg('Congratulations! Your path to awesomeness has began!')
```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 1e85212270
xp: 100
```

oho this is a new exercise ! mlkj mlkjlmj mlkj mlj klmjmlj lm!!

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 08658a03e4
xp: 50
```

`@projector_key`
d3480badef395e424cb56526b1627b1c
